Functionally-Coupled Ion Channels Begin Co-assembling at the Start of Their Synthesis

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public review):

Summary:

This manuscript by Pournejati et al investigates how BK (big potassium) channels and CaV1.3 (a subtype of voltage-gated calcium channels) become functionally coupled by exploring whether their ensembles form early-during synthesis and intracellular trafficking-rather than only after insertion into the plasma membrane. To this end, the authors use the PLA technique to assess the formation of ion channel associations in the different compartments (ER, Golgi or PM), single-molecule RNA in situ hybridization (RNAscope), and super-resolution microscopy.

Strengths:

The manuscript is well written and addresses an interesting question, combining a range of imaging techniques. The findings are generally well-presented and offer important insights into the spatial organization of ion channel complexes, both in heterologous and endogenous systems.

Weaknesses:

The authors have improved their manuscript after revisions, and some previous concerns have been addressed.

Still, the main concern about this work is that the current experiments do not quantitatively or mechanistically link the ensembles observed intracellularly (in the endoplasmic reticulum (ER) or Golgi) to those found at the plasma membrane (PM). As a result, it is difficult to fully integrate the findings into a coherent model of trafficking. Specifically, the manuscript does not address what proportion of ensembles detected at the PM originated in the ER. Without data on the turnover or halflife of these ensembles at the PM, it remains unclear how many persist through trafficking versus forming de novo at the membrane. The authors report the percentage of PLApositive ensembles localized to various compartments, but this only reflects the distribution of pre-formed ensembles. What remains unknown is the proportion of total BK and CaV1.3 channels (not just those in ensembles) that are engaged in these complexes within each compartment. Without this, it is difficult to determine whether ensembles form in the ER and are then trafficked to the PM, or if independent ensemble formation also occurs at the membrane. To support the model of intracellular assembly followed by coordinated trafficking, it would be important to quantify the fraction of the total channel population that exists as ensembles in each compartment. A comparable ensemble-to-total ratio across ER and PM would strengthen the argument for directed trafficking of pre-assembled channel complexes.

We appreciate the reviewer’s thoughtful comment and agree that quantitatively linking intracellular hetero-clusters to those at the plasma membrane is an important and unresolved question. Our current study does not determine what proportion of ensembles at the plasma membrane originated during trafficking. It also does not quantify the fraction of total BK and CaV1.3 channels engaged in these complexes within each compartment. Addressing this requires simultaneous measurement of multiple parameters—total BK channels, total CaV1.3 channels, hetero-cluster formation (via PLA), and compartment identity—in the same cell. This is technically challenging. The antibodies used for channel detection are also required for the proximity ligation assay, which makes these measurements incompatible within a single experiment.

To overcome these limitations, we are developing new genetically encoded tools to enable real-time tracking of BK and CaV1.3 dynamics in live cells. These approaches will enable us to monitor channel trafficking and the formation of hetero-clusters, as detected by colocalization. This kind of experiments will provide insight into their origin and turnover. While these experiments are beyond the scope of the current study, the findings in our current manuscript provide the first direct evidence that BK and CaV channels can form hetero-clusters intracellularly prior to reaching the plasma membrane. This mechanistic insight reveals a previously unrecognized step in channel organization and lays the foundation for future work aimed at quantifying ensemble-to-total ratios and determining whether coordinated trafficking of pre-assembled complexes occurs.

This limitation is acknowledged in the discussion section, page 23. It reads: “Our findings highlight the intracellular assembly of BK-CaV1.3 hetero-clusters, though limitations in resolution and organelle-specific analysis prevent precise quantification of the proportion of intracellular complexes that ultimately persist on the cell surface.”

Reviewer #2 (Public review):

Summary:

The co-localization of large conductance calcium- and voltage activated potassium (BK) channels with voltage-gated calcium channels (CaV) at the plasma membrane is important for the functional role of these channels in controlling cell excitability and physiology in a variety of systems.

An important question in the field is where and how do BK and CaV channels assemble as 'ensembles' to allow this coordinated regulation - is this through preassembly early in the biosynthetic pathway, during trafficking to the cell surface or once channels are integrated into the plasma membrane. These questions also have broader implications for assembly of other ion channel complexes

Using an imaging based approach, this paper addresses the spatial distribution of BKCaV ensembles using both overexpression strategies in tsa201 and INS-1 cells and analysis of endogenous channels in INS-1 cells using proximity ligation and superesolution approaches. In addition, the authors analyse the spatial distribution of mRNAs encoding BK and Cav1.3.

The key conclusion of the paper that BK and CaV1.3 are co-localised as ensembles intracellularly in the ER and Golgi is well supported by the evidence.However, whether they are preferentially co-translated at the ER, requires further work. Moreover, whether intracellular pre-assembly of BK-CaV1.3 complexes is the major mechanism for functional complexes at the plasma membrane in these models requires more definitive evidence including both refinement of analysis of current data as well as potentially additional experiments.

The reviewer raises the question of whether BK and CaV1.3 channels are preferentially co-translated. In fact, I would like to propose that co-translation has not yet been clearly defined for this type of interaction between ion channels. In our current work, we 1) observed the colocalization between BK and CaV1.3 mRNAs and 2) determined that 70% of BK mRNA in active translation also colocalizes with CaV1.3 mRNA. We think these results favor the idea of translational complexes that can underlie the process of co-translation. However, and in total agreement with the Reviewer, the conclusion that the mRNA for the two ion channels is cotranslated would require further experimentation. For instance, mRNA coregulation is one aspect that could help to define co-translation.

To avoid overinterpretation, we have revised the manuscript to remove references to “co-translation” in the Results section and included the word “potential” when referring to co-translation in the Discussion section. We also clarified the limitations of our evidence in the Discussion that can be found on page 25: “It is important to note that while our data suggest mRNA coordination, additional experiments are required to directly assess co-translation.”

Strengths & Weaknesses

(1) Using proximity ligation assays of overexpressed BK and CaV1.3 in tsa201 and INS1 cells the authors provide strong evidence that BK and CaV can exist as ensembles (ie channels within 40 nm) at both the plasma membrane and intracellular membranes, including ER and Golgi. They also provide evidence for endogenous ensemble assembly at the Golgi in INS-1 cells and it would have been useful to determine if endogenous complexes are also observe in the ER of INS-1 cells. There are some useful controls but the specificity of ensemble formation would be better determined using other transmembrane proteins rather than peripheral proteins (eg Golgi 58K).

We thank the reviewer for their thoughtful feedback and for recognizing the strength of our proximity ligation assay data supporting BK–CaV1.3 hetero-clusters formation at both the plasma membrane and intracellular compartments. As for specificity controls, we appreciate the suggestion to use transmembrane markers. To strengthen our conclusion, we have performed an additional experiment comparing the number of PLA puncta formed by the interaction of CaV1.3 and BK channels with the number of PLA puncta formed by the interaction of CaV1.3 channels and ryanodine receptors in INS-1 cells. As shown in the figure below, the number of interactions between CaV1.3 and BK channels is significantly higher than that between CaV1.3 and RyR2. Of note, RyR2 is a protein resident of the ER. These results provide additional evidence of the existence of endogenous complex formation in INS-1 cells. We have added this figure as a supplement.

(2) Ensemble assembly was also analysed using super-resolution (dSTORM) imaging in INS-1 cells. In these cells only 7.5% of BK and CaV particles (endogenous?) co-localise that was only marginally above chance based on scrambled images. More detailed quantification and validation of potential 'ensembles' needs to be made for example by exploring nearest neighbour characteristics (but see point 4 below) to define proportion of ensembles versus clusters of BK or Cav1.3 channels alone etc. For example, it is mentioned that a distribution of distances between BK and Cav is seen but data are not shown.

We thank the reviewer for this comment. To address the request for more detailed quantification and validation of ensembles, we performed additional analyses:

Proportion of ensembles vs isolated clusters: We quantified clusters within 200 nm and found that 37 ± 3% of BK clusters are near one or more CaV1.3 clusters, whereas 15 ± 2% of CaV1.3 clusters are near BK clusters. Figure 8– Supplementary 1A

Distance distribution: As shown in Figure 8–Supplementary 1B, the nearestneighbor distance distribution for BK-to-CaV1.3 in INS-1 cells (magenta) is shifted toward shorter distances compared to randomized controls (gray), supporting preferential localization of BK–CaV1.3 hetero-clusters.

Together, these analyses confirm that BK–CaV1.3 ensembles occur more frequently than expected by chance and exhibit an asymmetric organization favoring BK proximity to CaV1.3 in INS-1 cells. We have included these data and figures in the revised manuscript, as well as description in the Results section.

(3) The evidence that the intracellular ensemble formation is in large part driven by cotranslation, based on co-localisation of mRNAs using RNAscope, requires additional critical controls and analysis. The authors now include data of co-localised BK protein that is suggestive but does not show co-translation. Secondly, while they have improved the description of some controls mRNA co-localisation needs to be measured in both directions (eg BK - SCN9A as well as SCN9A to BK) especially if the mRNAs are expressed at very different levels. The relative expression levels need to be clearly defined in the paper. Authors also use a randomized image of BK mRNA to show specificity of co-localisation with Cav1.3 mRNA, however the mRNA distribution would not be expected to be random across the cell but constrained by ER morphology if cotranslated so using ER labelling as a mask would be useful?

We thank the reviewer for these constructive suggestions. We measured mRNA colocalization in both directions as recommended. As shown in the figure below, colocalization between KCNMA1 and SCN9A transcripts was comparable in both directions, with no statistically significant difference, supporting the specificity of the observed associations. We decided not to add this to the original figure to keep the figure simple.

We agree that co-localization of BK protein with BK mRNA is not conclusive evidence of co-translation, and we do not intend to mislead readers in our conclusion. Consequently, we were careful in avoiding the use of co-translation in the result section and added the word “potential” when referring to co-translation in the Discussion section. We added a sentence in the discussion to caution our interpretation: “It is important to note that while our data suggest mRNA coordination, additional experiments are required to directly assess cotranslation.”

(4) The authors attempt to define if plasma membrane assemblies of BK and CaV occur soon after synthesis. However, because the expression of BK and CaV occur at different times after transient transfection of plasmids more definitive experiments are required. For example, using inducible constructs to allow precise and synchronised timing of transcription. This would also provide critical evidence that co-assembly occurs very early in synthesis pathways - ie detecting complexes at ER before any complexes

We appreciate the reviewer’s insightful suggestion regarding the use of inducible constructs to synchronize transcription timing. This is an excellent approach and would allow direct testing of whether co-assembly occurs early in the synthesis pathway, including detection of complexes at the ER prior to plasma membrane localization. These experiments are beyond the scope of the present work but represent an important direction for future studies.

We have added the following sentence to the Discussion section (page 24) to highlight this idea. “Future experiments using inducible constructs to precisely control transcription timing will enable more precise quantification of heterocluster formation in the ER compartment prior to plasma membrane insertion and reduce the variability introduced by differences in expression timing after plasmid transfection.”

(5) While the authors have improved the definition of hetero-clusters etc it is still not clear in superesolution analysis, how they separate a BK tetramer from a cluster of BK tetramers with the monoclonal antibody employed ie each BK channel will have 4 binding sites (4 subunits in tetramer) whereas Cav1.3 has one binding site per channel. Thus, how do authors discriminate between a single BK tetramer (molecular cluster) with potential 4 antibodies bound compared to a cluster of 4 independent BK channels.

We appreciate the reviewer’s thoughtful comment regarding the interpretation of super-resolution data. We agree that distinguishing a single BK tetramer from a cluster of multiple BK channels is challenging when using an antibody that can bind up to four sites per channel. To clarify, our analysis does not attempt to resolve individual subunits within a tetramer; rather, it focuses on the nanoscale spatial proximity of BK and CaV1.3 signals.

We want to note that this limitation applies only to the super-resolution maps in Figures 8C and 9D and does not affect Airyscan-based analyses or measurements of BK–CaV1.3 proximity.

To address how we might distinguish between a single BK tetramer and a cluster of multiple BK channels, we considered two contrasting scenarios. In the first case, we assume that all four α-subunits within a tetramer are labeled. Based on cryoEM structures, a BK tetramer measures approximately 13 nm × 13 nm (≈169 nm²). Adding two antibody layers (primary and secondary) would increase the footprint by ~14 nm in each direction, resulting in an estimated area of ~41 nm × 41 nm (≈1681 nm²). Under this assumption, particles smaller than ~1681 nm² would likely represent individual tetramers, whereas larger particles would correspond to clusters of multiple tetramers.

In the second scenario, we propose that steric constraints at the S9–S10 segment, where the antibody binds, limit labeling to a single antibody per tetramer. If true, the localization precision would approximate 14 nm × 14 nm—the combined size of the antibody complex and the channel—close to the resolution limit of the microscope. To test this, we performed a control experiment using two antibodies targeting the BK C-terminal domain, raised in different species and labeled with distinct fluorophores. Super-resolution imaging revealed that only ~12% of particles were colocalized, suggesting that most channels bind a single antibody.

If multiple antibodies could bind each tetramer, we would expect much greater colocalization.

Although these data are not included in the manuscript, we have added the following clarification to the Results section (page 19): “It is important to note that this technique does not allow us to distinguish between labeling of four BK αsubunits within a tetramer and labeling of multiple BK channel clusters. Hence, particles smaller than ~1680 nm² may represent either a single tetramer or a cluster. This limitation applies to Figures 8C and 9D and does not affect measurements of BK–CaV1.3 proximity.”

(6) The post-hoc tests used for one way ANOVA and ANOVA statistics need to be defined throughout

We thank the reviewer for highlighting the need for clarity regarding our statistical analyses. We have now specified the post-hoc tests used for all one-way ANOVA and ANOVA comparisons throughout the manuscript, and updated figure legends.

Reviewer #3 (Public review):

Summary:

The authors present a clearly written and beautifully presented piece of work demonstrating clear evidence to support the idea that BK channels and Cav1.3 channels can co-assemble prior to their assertion in the plasma membrane.

Strengths:

The experimental records shown back up their hypotheses and the authors are to be congratulated for the large number of control experiments shown in the ms.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

The authors have sufficiently addressed the specific points previously raised and the manuscript has improved clarity in those aspects. My main concern, which still remains, is stated in the public review.

Reviewer #3 (Recommendations for the authors):

I am content that the authors have attempted to fully address my previous criticisms.

I have only three suggestions

(1) I think the word Homo-clusters at the bottom right of Figure 1 is erroneously included.

We thank the reviewer for bringing this to our attention. The figure has been corrected accordingly.

(2) The authors should, for completeness, to refer to the beta, gamma and LINGO subunit families in the Introduction and include appropriate references:

Knaus, H. G., Folander, K., Garcia-Calvo, M., Garcia, M. L., Kaczorowski, G. J., Smith, M., & Swanson, R. (1994). Primary sequence and immunological characterization of betasubunit of high conductance Ca2+-activated K+ channel from smooth muscle. The Journal of Biological Chemistry, 269(25), 17274-17278.

Brenner, R., Jegla, T. J., Wickenden, A., Liu, Y., & Aldrich, R. W. (2000a). Cloning and functional characterization of novel large conductance calcium-activated potassium channel beta subunits, hKCNMB3 and hKCNMB4. The Journal of Biological Chemistry, 275(9), 6453-6461.

Yan, J & R.W. Aldrich. (2010) LRRC26 auxiliary protein allows BK channel activation at resting voltage without calcium. Nature. 466(7305):513-516

Yan, J & R.W. Aldrich. (2012) BK potassium channel modulation by leucine-rich repeatcontaining proteins. Proceedings of the National Academy of Sciences 109(20):7917-22

Dudem, S, Large RJ, Kulkarni S, McClafferty H, Tikhonova IG, Sergeant, GP, Thornbury, KD, Shipston, MJ, Perrino BA & Hollywood MA (2020). LINGO1 is a novel regulatory subunit of large conductance, Ca2+-activated potassium channels. Proceedings of the National Academy of Sciences 117 (4) 2194-2200

Dudem, S., Boon, P. X., Mullins, N., McClafferty, H., Shipston, M. J., Wilkinson, R. D. A., Lobb, I., Sergeant, G. P., Thornbury, K. D., Tikhonova, I. G., & Hollywood, M. A. (2023). Oxidation modulates LINGO2-induced inactivation of large conductance, Ca2+-activated potassium channels. The Journal of Biological Chemistry, 299 (3) 102975.

We agree with the reviewer’s suggestion and have revised the Introduction to include references to the beta, gamma, and LINGO subunit families. Appropriate citations have been added to ensure completeness and contextual relevance.

Additionally, BK channels are modulated by auxiliary subunits, which fine-tune BK channel gating properties to adapt to different physiological conditions. The β, γ, and LINGO1 subunits each contribute distinct structural and regulatory features: β-subunits modulate Ca²⁺ sensitivity and can induce inactivation; γ-subunits shift voltage-dependent activation to more negative potentials; and LINGO1 reduces surface expression and promotes rapid inactivation (18-24). These interactions ensure precise control over channel activity, allowing BK channels to integrate voltage and calcium signals dynamically in various cell types.

(3) I think it may be more appropriate to include the sentence "The probes against the mRNAs of interest and tested in this work were designed by Advanced Cell Diagnostics." (P16, right hand column, L12-14) in the appropriate section of the Methods, rather than in Results.

We thank the reviewer for this helpful suggestion. In response, we have relocated the sentence to the appropriate section of the Methods, where it now appears with relevant context.

Changes from the earlier version to the revised version

Abstract

Calcium binding to BK channels lowers BK activation threshold, substantiating functional coupling with calcium-permeable channels. This coupling requires close proximity between different channel types, and the formation of BK–CaV1.3 hetero-clusters at nanometer distances exemplifies this unique organization. To investigate the structural basis of this interaction, we tested the hypothesis that BK and CaV1.3 channels assemble before their insertion into the plasma membrane. Our approach incorporated four strategies: (1) detecting interactions between BK and CaV1.3 proteins inside the cell, (2) identifying membrane compartments where intracellular hetero-clusters reside, (3) measuring the proximity of their mRNAs, and (4) assessing protein interactions at the plasma membrane during early translation. These analyses revealed that a subset of BK and CaV1.3 transcripts are spatially close in micro-translational complexes, and their newly synthesized proteins associate within the endoplasmic reticulum (ER) and Golgi. Comparisons with other proteins, transcripts, and randomized localization models support the conclusion that BK and CaV1.3 hetero-clusters form before their insertion at the plasma membrane.

Introduction

This work concerns the organization of functionally coupled voltage- and calcium-dependent potassium (BK) channels and voltage-gated calcium (CaV) channels.

BK channels (a.k.a. KCa1.1, Maxi-K, Slo1, KCNMA1) are named for their “big potassium (K+)” conductance, since channel opening leads to unusually large, outward potassium currents and membrane repolarization. Its amino acid sequence is conserved throughout the animal kingdom from worms to mammals. BK channels are expressed in a wide variety of cell types, predominantly excitable cells but also non-excitable salivary, bone and kidney cells, making this channel responsible for a large range of physiological processes [1–9]. Phenotypes of pathological BK mutations in human patients are most prominent in the brain and muscle and often manifest as seizures, movement disorders, developmental delay, and intellectual disability [10, 11]. BK is also implicated in other organ system functions, including but not limited to heart pace making, reproduction, and pancreatic glucose homeostasis [11, 12].

Since BK channels affect function in numerous physiological systems, channel opening is meticulously regulated at the cellular level. BK is a voltage-gated channel activated by membrane depolarization. Interestingly, BK has an additional gating mechanism that differentiates it from typical voltage-gated potassium channels; BK opening is gated by both voltage and intracellular calcium binding to a cytosolic regulatory domain. At the cytoplasmic resting free calcium concentration (∼0.1 µM), BK channels remain closed [13]. However, the probability of channel opening increases when both the membrane potential depolarizes and when the local free calcium rises. Yet, increases of calcium are tightly limited by endogenous proteins that bind calcium with high affinity as well as by extrusion mechanisms that take calcium inside organelles or outside the cell. Hence, the activation of BK channels relies on strategies to overcome these regulatory barriers. One such strategy is to localize near sources of calcium.

Earlier version: In excitable cells, BK forms nanodomains with calcium channels that provide exclusive, localized calcium sources. Several subtypes of calcium channels form functional signaling complexes with BK, shifting BK activation voltages to more negative potentials [14-16]. One of these channels, CaV1.3 (a.k.a. CACNA1D), is unique in its electrophysiological profile as an L-type calcium channel, activating with fast kinetics at voltages as negative as −55 mV [17]. Additionally, BK channels are modulated by auxiliary subunits, which fine-tune BK channel gating properties to adapt to different physiological conditions. β and γ subunits regulate BK channel kinetics, altering voltage sensitivity and calcium responsiveness [18]. These interactions ensure precise control over channel activity, allowing BK channels to integrate voltage and calcium signals dynamically in various cell types. Here, we focus on the selective assembly of BK channels with CaV1.3 and do not evaluate the contributions of auxiliary subunits to BK channel organization. CaV1.3 is expressed in many of the same cell types as BK and is often functionally coupled with BK channels, enabling BK channels to activate at more negative voltages. Notably, super-resolution microscopy shows that CaV1.3 organizes spatially into nanodomains with BK in the plasma membrane [15]. However, the mechanisms behind the assembly of BK and CaV1.3 hetero-clusters are unknown.

Revised version: In excitable cells, BK forms nanodomains with calcium channels that provide exclusive, localized calcium sources. Several subtypes of calcium channels form functional signaling complexes with BK, shifting BK activation voltages to more negative potentials [14–16]. One of these channels, CaV1.3 (a.k.a. CACNA1D), is unique in its electrophysiological profile as an L-type calcium channel, activating with fast kinetics at voltages as negative as −55 mV [17]. Additionally, BK channels are modulated by auxiliary subunits, which fine-tune BK channel gating properties to adapt to different physiological conditions. The β, γ, and LINGO1 subunits each contribute distinct structural and regulatory features: β-subunits modulate Ca²⁺ sensitivity and can induce inactivation; γ-subunits shift voltage-dependent activation to more negative potentials; and LINGO1 reduces surface expression and promotes rapid inactivation. These interactions ensure precise control over channel activity, allowing BK channels to integrate voltage and calcium signals dynamically in various cell types [18–24].

Added in the revised version: Here, we focus on the selective assembly of BK channels with CaV1.3 and do not evaluate the contributions of auxiliary subunits to BK channel organization. CaV1.3 is expressed in many of the same cell types as BK and is often functionally coupled with BK channels, enabling BK channels to activate at more negative voltages. Notably, super-resolution microscopy shows that CaV1.3 organizes spatially into nanodomains with BK in the plasma membrane [15]. However, the mechanisms behind the assembly of BK and CaV1.3 hetero-clusters are unknown.

Earlier version: Several mechanisms for bringing ion channels together have been suggested (Table 1). One proposes that proteins assembly precedes protein insertion into the plasma membrane [19, 25, 26]. This mechanism has been observed in studies of hetero-multimers and even channels permeating different ions. The formation of these groups of proteins would start from their synthesis. There is even the possibility of mRNA transcripts colocalizing prior to translation. We explored this mechanism in relation to BK and CaV1.3 functional coupling. Here, we started by investigating when and where BK and CaV1.3 channels cluster in the cell. We looked for BK and CaV1.3 hetero-clusters at intracellular membranes of the ER, at ER exit sites, and at the Golgi. We also investigated the proximity between mRNA transcripts of BK and CaV1.3.

Revised version: Several mechanisms for bringing ion channels together have been suggested (Table 1). One proposes that proteins assembly precedes protein insertion into the plasma membrane [25, 26, 31]. This mechanism has been observed in studies of hetero-multimers and even channels permeating different ions. The formation of these groups of proteins would start from their synthesis. There is even the possibility of mRNA transcripts colocalizing prior to translation. We explored this mechanism in relation to BK and CaV1.3 functional coupling. Here, we started by investigating when and where BK and CaV1.3 channels cluster in the cell. We looked for BK and CaV1.3 hetero-clusters at intracellular membranes of the ER, at ER exit sites, and at the Golgi. We also investigated the proximity between mRNA transcripts of BK and CaV1.3.

Definitions used in this manuscript

To guide the reader and prevent confusion on the terms used, we introduce the following nomenclature, which is also illustrated in Figure 1. Molecular complex: an array of several polypeptides with a defined function. In our case, a BK channel is a molecular complex of four alpha subunits whose function is to permeate ions. Homo-cluster: we define a homo-cluster as the accumulation of proteins. In our work, a homo-cluster refers to the accumulation of BK channels or the accumulation of calcium channels.

Representation of molecular complex, homo-cluster, and hetero-cluster.

Hetero-cluster: we define a hetero-cluster as the collection of different types of proteins that facilitate functional coupling and compartmentalization of a signaling complex. In the present work, hetero-cluster refers to a coordinated collection of BK and CaV channels. It is also used to refer to the assembly of homo-clusters of BK with homo-clusters of CaV channels.

Materials and Methods

Cell culture

We used tsA-201 cells to co-express BK and CaV1.3 channels heterologously. Cells were grown in DMEM (Gibco) supplemented with 10% fetal bovine serum and 0.2% penicillin/streptomycin. We used rat insulinoma (INS-1) cells to study endogenous levels of transcripts and proteins of channels. INS-1 cells were cultured in RPMI high glutamate medium (Gibco) with 10% fetal bovine serum, 0.2% penicillin/streptomycin, 10 mM HEPES (Gibco), 1 mM sodium pyruvate (Gibco), and 50 μM 2-mercaptoethanol. Both cell types were passaged twice a week and incubated in 5% CO2 at 37°C.

Plasmids and transfection

Cells were transfected with 0.1– 0.4 μg DNA per plasmid and plated for 24 hours on poly-D-Lysine coated coverslips. Lipofectamine 3000 (Invitrogen, RRID: L30000) was used for the transfection. DNA clones of CaV1.3, BK channels, PH-PLCδ GFP, ER moxGFP, pmGFP-Sec16S and Golgi-mGFP were obtained from Addgene (RRID:SCR_002037). The CaV1.3 α subunit construct used in our study corresponds to the rat CaV1.3e splice variant containing exons 8a, 11, 31b, and 42a, with a deletion of exon 32. The BK channel construct corresponds to the VYR splice variant of the mouse BKα subunit (KCNMA1). Auxiliary subunits for CaV1.3 channels, CaVβ3 and CaVα2δ1 (from Diane Lipscombe, Brown University, RI), were transfected as well. No BK channel auxiliary subunits were transfected.

Antibodies

Earlier version: CaV1.3 channels were immuno-detected with a rabbit primary antibody recognizing residues 809 to 825 located at the intracellular II-III loop of the channel (DNKVTIDDYQEEAEDKD), kindly provided by Drs. William Catterall and Ruth Westenbroek [27]. BK channels were detected using the anti-Slo1 mouse monoclonal antibody clone L6/60. The goat polyclonal GFP antibody was against the recombinant full-length protein corresponding to Aequorea victoria GFP. Anti-58K Golgi protein antibody was used to mark the Golgi. Specificity of antibodies was tested in un-transfected tsA-201 cells (Figure 2-Supplementary 1). The secondary antibodies tagged with Alexa dyes were Donkey anti-mouse Alexa-647, Donkey anti-rabbit Alexa-555, Donkey anti-goat Alexa-488 (Molecular Probes).

Revised version: CaV1.3 channels were immuno-detected with a rabbit primary antibody recognizing residues 809 to 825 located at the intracellular II-III loop of the channel (DNKVTIDDYQEEAEDKD), kindly provided by Drs. William Catterall and Ruth Westenbroek [32]. BK channels were detected using the anti-Slo1 mouse monoclonal antibody clone L6/60. The goat polyclonal GFP antibody was against the recombinant full-length protein corresponding to Aequorea victoria GFP. Anti-58K Golgi protein antibody was used to mark the Golgi. Specificity of antibodies was tested in un-transfected tsA-201 cells (Figure 2-Supplementary 1). The secondary antibodies tagged with Alexa dyes were Donkey anti-mouse Alexa-647, Donkey anti-rabbit Alexa-555, Donkey anti-goat Alexa-488 (Molecular Probes).

BK and CaV1.3 hetero-clusters are found inside the cell.

Earlier version: A. Diagram of the hypothesis: hetero-clusters of BK (magenta) and CaV1.3 (cyan) are on intracellular membranes and on the plasma membrane. B. Illustration of the technique to detect BK and CaV1.3 hetero-clusters. Proximity ligation assay is used to detect the hetero-clusters. C. Confocal images of fluorescent puncta from PLA experiments in tsA-201 cells. Left: Cells were transfected and probed for BK and CaV1.3 channels. Right: negative control. Cells were transfected and probed only for BK channels. Enlargement of a selected region is shown in the inset. D. Scatter dot plot comparing puncta density of BK and CaV1.3 hetero-clusters to the negative control. Data points are from n = 12 cells for BK and CaV1.3 hetero-clusters and from n = 14 cells for negative control. p-values are shown at the top of the graphs. E. Confocal images of fluorescent PLA puncta at different focal planes co-labeled against GFP at the plasma membrane. Cells were transfected with BK, CaV1.3, and PH-PLCδ-GFP and probed for BK channels, CaV1.3 channels, and GFP. PLA puncta are shown in magenta, and the plasma membrane is shown in green. Enlargements of the representative regions of PM and intercellular hetero-clusters are shown in the insets. F. Scatter dot plot comparing BK and CaV1.3 hetero-cluster abundance at PM and inside the cell. Data points are from n = 12 cells. Scale bar is 10 μm and 1 μm in the insets.

Revised version: A. Diagram of the hypothesis: hetero-clusters of BK (magenta) and CaV1.3 (cyan) are on intracellular membranes and on the plasma membrane. B. Illustration of the technique to detect BK and CaV1.3 hetero-clusters. Proximity ligation assay is used to detect the hetero-clusters. C. Confocal images of fluorescent puncta from PLA experiments in tsA-201 cells. Left: Cells were transfected and probed for BK and CaV1.3 channels. Right: negative control. Cells were transfected and probed only for BK channels. Enlargement of a selected region is shown in the inset. D. Scatter dot plot comparing puncta density of BK and CaV1.3 hetero-clusters to the negative control. Data points are from n = 12 cells for BK and CaV1.3 hetero-clusters and from n = 14 cells for negative control. p-values are shown at the top of the graphs. E. Confocal images of fluorescent PLA puncta at different focal planes co-labeled against GFP at the plasma membrane. Cells were transfected with BK, CaV1.3, and PH-PLCδ-GFP and probed for BK channels, CaV1.3 channels, and GFP. PLA puncta are shown in magenta, and the plasma membrane is shown in green. Enlargements of the representative regions of PM and intercellular hetero-clusters are shown in the insets. F. Scatter dot plot comparing BK and CaV1.3 hetero-cluster abundance at PM and inside the cell. Data points are from n = 12 cells. Scale bar is 10 μm and 1 μm in the insets. For statistical analysis, unpaired t-tests were applied in panel D and paired t-tests in panel F to evaluate significance.

Removed in the revised version: Validation of antibodies against BK, CaV1.3 and GFP.

Removed in the revised version: Representative confocal images of tsA-201 cells immuno-tested for BK channels (left), CaV1.3 channels (middle), and GFP proteins (right). Cells were not transfected. Nuclei were stained with DAPI and pseudo colored in gray. Scale bar is 10 μm for all images.

Immunostaining

Cells were fixed with freshly prepared 4% paraformaldehyde for 10 minutes. After washing, aldehydes were reduced with 0.1% NaBH4 for 5 minutes and then washed again. Nonspecific binding was blocked with 3% bovine serum albumin (Thermo Scientific). Cells were permeabilized with 0.25% v/v Triton X-100 in PBS for 1 hour. Primary antibodies were used at 10 μg/ml in blocking solution and incubated overnight at 4°C. After washing, secondary antibodies at 2 μg/ml were incubated for 1 hour at 21°C. Washing steps indicated in all methods include 3 cycles of rinsing and rocking for 5 minutes with PBS at 21°C. Cells were imaged using an inverted AiryScan microscope or an ONI Nanoimager with super-resolution capabilities, in total internal reflection fluorescence (TIRF) mode, and with a Z-resolution of 50 nm.

Proximity ligation assay (PLA)

Cells were fixed with freshly prepared 4% paraformaldehyde for 10 minutes. After washing, aldehydes were reduced with 50 mM glycine for 15 minutes. After another round of washes, PLA was performed according to manufacturer instructions (Duolink® In Situ Red Starter Kit). Cells were blocked and permeabilized with Duolink blocking solution. Primary antibodies were used at 10 μg/ml in Duolink antibody diluent and incubated overnight at 4°C. The Duolink® In Situ PLA® probe anti-rabbit PLUS and anti-mouse MINUS were used as secondary antibodies, followed by ligation and amplification. For PLA combined with immunostaining, PLA was followed by a secondary antibody incubation with Alexa Fluor-488 at 2 μg/ml for 1 hour at 21°C. Since GFP fluorescence fades significantly during the PLA protocol, resulting in reduced signal intensity and poor image resolution, GFP was labeled using an antibody rather than relying on its intrinsic fluorescence. Coverslips were mounted using ProLong™ Gold Antifade Mountant with DAPI. Cells were imaged using an inverted Zeiss AiryScan microscope.

Earlier version: Single-molecule fluorescence in situ hybridization (RNAscope™)

Revised version: Single-molecule fluorescence in situ hybridization (RNAscopeTM)

Earlier version: Manual RNAscope assay was performed using RNAscope Multiplex Fluorescent V2 Assay according to the manufacturer’s protocol. The RNAscope assay consists of target probes and a signal amplification system composed of a preamplifier, amplifier, and label probe. A schematic RNAscope assay procedure is shown in Figure 5-Supplementary 1. Briefly, cells were fixed with 4% paraformaldehyde for 30 minutes, washed, dehydrated, and then rehydrated with ethanol, and permeabilized with 0.1% Tween-20 in PBS. Next, cells were quenched with H2O2 and treated with Protease III. Probes were hybridized for 2 hours at 40° C followed by RNAscope amplification and then fluorescence detection. Coverslips were mounted using ProLong™ Gold Antifade Mountant with DAPI. We used the following RNAscope probes: RNAscope 3-plex Positive Control Probes, RNAscope 3-plex negative control probes, RNAscope™ Probe-Rn-Ryr, RNAscope™ Probe-Rn-Scn9a, RNAscope Probe-Rn-Kcnma1-C3, RNAscope Probe-Rn-Cacna1d-C2, and RNAscope ProbeRn-Gapdh. Cells were imaged on the inverted AiryScan microscope. For PLA and RNAscope experiments, we used custom-made macros written in ImageJ. Processing of PLA data included background subtraction. To assess colocalization, fluorescent signals were converted into binary images, and channels were multiplied to identify spatial overlap. Specificity of RNAscope probes was tested in un-transfected tsA-201 cells (Figure 5-Supplementary 2). For RNAscope combined with immunostaining, RNAscope was followed by blocking in PBS supplemented with 0.01 % Tween-20 and 3% BSA for 1 h at 21°C. Samples were then probed for BK protein using primary antibody overnight at 4°C followed by secondary antibody incubation with Alexa Fluor-488 at 2 μg/ml for 1 hour at 21°C. Coverslips were mounted using ProLong™ Gold Antifade Mountant. Cells were imaged using an inverted Zeiss AiryScan microscope.

Revised version: Manual RNAscope assay was performed using RNAscope Multiplex Fluorescent V2 Assay according to the manufacturer’s protocol. The RNAscope assay consists of target probes and a signal amplification system composed of a preamplifier, amplifier, and label probe. A schematic RNAscope assay procedure is shown in Figure 5-Supplementary 1. Briefly, cells were fixed with 4% paraformaldehyde for 30 minutes, washed, dehydrated, and then rehydrated with ethanol, and permeabilized with 0.1% Tween-20 in PBS. Next, cells were quenched with H2O2 and treated with Protease III. Probes were hybridized for 2 hours at 40°C followed by RNAscope amplification and then fluorescence detection. Coverslips were mounted using ProLong™ Gold Antifade Mountant with DAPI. We used the following RNAscope probes: RNAscope 3-plex Positive Control Probes, RNAscope 3-plex negative control probes, RNAscope™ Probe- Rn-Ryr, RNAscope™ Probe- Rn-Scn9a, RNAscope Probe- Rn-Kcnma1-C3, RNAscope Probe- Rn-Cacna1d-C2, and RNAscope Probe- Rn-Gapdh. Cells were imaged on the inverted AiryScan microscope. For PLA and RNAscope experiments, we used custom-made macros written in ImageJ. Processing of PLA data included background subtraction. To assess colocalization, fluorescent signals were converted into binary images, and channels were multiplied to identify spatial overlap. Specificity of RNAscope probes was tested in un-transfected tsA-201 cells (Figure 5-Supplementary 2). For RNAscope combined with immunostaining, RNAscope was followed by blocking in PBS supplemented with 0.01 % Tween-20 and 3% BSA for 1 h at 21°C. Samples were then probed for BK protein using primary antibody overnight at 4°C followed by secondary antibody incubation with Alexa Fluor-488 at 2 μg/ml for 1 hour at 21°C. Coverslips were mounted using ProLong™ Gold Antifade Mountant. Cells were imaged using an inverted Zeiss AiryScan microscope.

High resolution imaging

Cells were imaged using an inverted AiryScan microscope (Zeiss LSM 880) run by ZEN black v2.3 software and equipped with a plan apochromat 63X oil immersion objective with 1.4 NA. Fluorescent dyes were excited with a 405 nm diode, 458–514 nm argon, 561 nm, or 633 nm laser. Emission light was detected using an Airyscan 32 GaAsP detector and appropriate emission filter sets. The point spread functions were calculated using ZEN black software and 0.1 μm fluorescent microspheres. The temperature inside the microscope housing was 22°C. Images were analyzed using ImageJ (NIH).

Super-resolution imaging

Direct stochastic optical reconstruction microscopy (dSTORM) images of BK and CaV1.3 overexpressed in tsA-201 cells were acquired using an ONI Nanoimager microscope equipped with a 100X oil immersion objective (1.4 NA), an XYZ closed-loop piezo 736 stage, and triple emission channels split at 488, 555, and 640 nm. Samples were imaged at 35°C. For single-molecule localization microscopy, fixed and stained cells were imaged in GLOX imaging buffer containing 10 mM β-mercaptoethylamine (MEA), 0.56 mg/ml glucose oxidase, 34 μg/ml catalase, and 10% w/v glucose in Tris-HCl buffer. Single-molecule localizations were filtered using NImOS software (v.1.18.3, ONI). Localization maps were exported as TIFF images with a pixel size of 5 nm. Maps were further processed in ImageJ (NIH) by thresholding and binarization to isolate labeled structures. To assess colocalization between the signal from two proteins, binary images were multiplied. Particles smaller than 400 nm2 were excluded from the analysis to reflect the spatial resolution limit of STORM imaging (20 nm) and the average size of BK channels. To examine spatial localization preference, binary images of BK were progressively dilated to 20 nm, 40 nm, 60 nm, 80 nm, 100 nm, and 200 nm to expand their spatial representation. These modified images were then multiplied with the CaV1.3 channel to quantify colocalization and determine BK occupancy at increasing distances from CaV1.3. To ensure consistent comparisons across distance thresholds, data were normalized using the 200 nm measurement as the highest reference value, set to 1.

Image scrambling

Earlier version: Images were binarized as TIFF images, and their respective cell perimeter coordinates were exported as CSV files by ImageJ. Processed binary images were then analyzed by our SpotScrambler (https://github.com/jehuang2/SpotScrambler) Python program. SpotScrambler first extracts the areas of fluorescent particles in the binary image. SpotScrambler then redraws the fluorescent particles as circles at randomized coordinates within cell perimeter boundaries. SpotScrambler accurately preserves particle number and particle sizes, averaging less than 1% difference in total area of fluorescent particles between pre-SpotScrambler and post-SpotScrambler images. To ensure reliable randomization for each experiment, results were averaged between 3 trials of SpotScrambler.

Revised version: Images were binarized as TIFF images, and their respective cell perimeter coordinates were exported as CSV files by ImageJ. Processed binary images were then analyzed by our SpotScrambler (https://github.com/jehuang2/SpotScramble) Python program. SpotScrambler first extracts the areas of fluorescent particles in the binary image. SpotScrambler then redraws the fluorescent particles as circles at randomized coordinates within cell perimeter boundaries. SpotScrambler accurately preserves particle number and particle sizes, averaging less than 1% difference in total area of fluorescent particles between pre-SpotScrambler and post-SpotScrambler images. To ensure reliable randomization for each experiment, results were averaged between 3 trials of SpotScrambler.

Data analysis

Excel (Microsoft), and Prism (GraphPad) were used to analyze data. ImageJ was used to process images. One-way ANOVA and non-parametric statistical test (Mann-Whitney Wilcoxon) were used to test for statistical significance. p-values <0.05 were deemed statistically significant. The number of cells used for each experiment is detailed in each figure legend.

Results

BK and CaV1.3 hetero-clusters are found inside the cell

BK and CaV hetero-clusters have been observed at the plasma membrane [14–16]. However, a clear understanding of when or how BK and CaV hetero-clusters assemble is missing. A simple mechanism proposes that these hetero-clusters organize only at the plasma membrane. Here, we tested the alternative hypothesis of BK and CaV assembling inside the cell (Figure 2A). To test this idea, we used proximity ligation assay (PLA) and antibodies against BK and CaV1.3 channels. When these antibodies are within 40 nm of each other, PLA ligation and amplification can occur, resulting in the formation of fluorescent puncta (Figure 2B), here referred to as PLA puncta. Hence, the PLA puncta in Figure 2C represent BK and CaV1.3 hetero-clusters. To confirm specificity, a negative control was performed by probing only for BK using the primary antibody, ensuring that detected signals were not due to non-specific binding or background fluorescence. We first analyzed Z-projections, in which all the puncta in the cell volume are added, and found a density of 5.8 ± 1.0 / 10 μm2. In a different experiment, we analyzed the puncta density for each focal plane of the cell (step size of 300 nm) and compared the puncta at the plasma membrane to the rest of the cell. We visualized the plasma membrane with a biological sensor tagged with GFP (PH-PLCδ-GFP) and then probed it with an antibody against GFP (Figure 2E). By analyzing the GFP signal, we created a mask that represented the plasma membrane. The mask served to distinguish between the PLA puncta located inside the cell and those at the plasma membrane, allowing us to calculate the number of PLA puncta at the plasma membrane. To our surprise, we found a significant number of puncta localized inside the cell. 46 ± 3 % of the puncta were localized intracellularly, whereas 54 ± 3 % were at the plasma membrane. This finding is consistent with our supposition that BK and CaV1.3 channels colocalize in the cell.

BK and CaV1.3 hetero-clusters localize at ER and ER exit sites

Earlier version: We next investigated the identity of the intracellular membranes where these PLA puncta were found. A large component of intracellular membranes in the cell is the endoplasmic reticulum (ER), where channels are inserted after translation. To determine whether BK and CaV1.3 associate in the ER (Figure 3A), we combined PLA with immunodetection. As before, PLA probed for BK and CaV1.3 hetero-clusters, while a KDEL-moxGFP label identified the ER (Figure 3B). To avoid disruption of organelle architecture, we used the monomeric mox version of KDEL-GFP, which is optimized to reduce oligomerization in the ER environment [28]. Neither overexpression of KDEL-moxGFP nor fixation altered the ER structure as the tubule width remained around 150 nm for either condition (Figure 3C), a value in agreement with the literature [29, 30]. To assess the percentage of PLA puncta colocalizing with the ER, we employed two different cell lines, our overexpression system (tsA-201 cells) and a rat insulinoma cell line (INS-1) that expresses BK and CaV1.3 channels endogenously. In this and following experiments, we analyze one focal plane, in the middle of the cell, to quantify PLA puncta colocalization with ER membrane. Figure 3D shows PLA puncta in the same space as the ER. Comparing the overexpression and endogenous systems, 63 ± 3% versus 50 ± 6% of total PLA puncta were localized at the ER (Figure 3E). To determine whether the observed colocalization between BK–CaV1.3 hetero-clusters and the ER was not simply due to the extensive spatial coverage of ER labeling, we labeled ER exit sites using Sec16-GFP and probed for hetero-clusters with PLA. This approach enabled us to test whether the hetero-clusters were preferentially localized to ER exit sites, which are specialized trafficking hubs that mediate cargo selection and direct proteins from the ER into the secretory pathway. In contrast to the more expansive ER network, which supports protein synthesis and folding, ER exit sites ensure efficient and selective export of proteins to their target destinations.

Revised version: We next investigated the identity of the intracellular membranes where these PLA puncta were found. A large component of intracellular membranes in the cell is the endoplasmic reticulum (ER), where channels are inserted after translation. To determine whether BK and CaV1.3 associate in the ER (Figure 3A), we combined PLA with immunodetection. As before, PLA probed for BK and CaV1.3 hetero-clusters, while a KDEL-moxGFP label identified the ER (Figure 3B). To avoid disruption of organelle architecture, we used the monomeric mox version of KDEL-GFP, which is optimized to reduce oligomerization in the ER environment [33]. Neither overexpression of KDEL-moxGFP nor fixation altered the ER structure as the tubule width remained around 150 nm for either condition (Figure 3C), a value in agreement with the literature [34, 35]. To assess the percentage of PLA puncta colocalizing with the ER, we employed two different cell lines, our overexpression system (tsA-201 cells) and a rat insulinoma cell line (INS-1) that expresses BK and CaV1.3 channels endogenously. In this and following experiments, we analyze one focal plane, in the middle of the cell, to quantify PLA puncta colocalization with ER membrane. Figure 3D shows PLA puncta in the same space as the ER. Comparing the overexpression and endogenous systems, 63 ± 3% versus 50 ± 6% of total PLA puncta were localized at the ER (Figure 3E). To determine whether the observed colocalization between BK–CaV1.3 hetero-clusters and the ER was not simply due to the extensive spatial coverage of ER labeling, we labeled ER exit sites using Sec16-GFP and probed for hetero-clusters with PLA. This approach enabled us to test whether the hetero-clusters were preferentially localized to ER exit sites, which are specialized trafficking hubs that mediate cargo selection and direct proteins from the ER into the secretory pathway. In contrast to the more expansive ER network, which supports protein synthesis and folding, ER exit sites ensure efficient and selective export of proteins to their target destinations. By quantifying the proportion of BK and CaV1.3 hetero-clusters relative to total channel expression at ER exit sites, we found 28 ± 3% colocalization in tsA-201 cells and 11 ± 2% in INS-1 cells (Figure 3F). While the percentage of colocalization between hetero-clusters and the ER or ER exit sites alone cannot be directly compared to infer trafficking dynamics, these findings reinforce the conclusion that hetero-clusters reside within the ER and suggest that BK and CaV1.3 channels traffic together through the ER and exit in coordination.

BK and CaV1.3 hetero-clusters localize at ER and ER exit sites (ERES).

A. Diagram of the hypothesis: hetero-clusters of BK (magenta) and CaV1.3 (cyan) can be found at the ER membrane. B. Representative image of the ER labeled with exogenous GFP in INS-1 cells. Cells were transfected with KDEL-moxGFP. Magnification is shown in the inset. C. Comparison of the ER tubule distance in live and fixed tsA-201 and INS-1 cells. Data points are from n = 23 tsA-201 cells, n = 27 INS-1 cells D. Representative images of PLA puncta and ER. Left: tsA-201 cells were transfected with BK, CaV1.3 and KDEL-moxGFP. Right: INS-1 cells were transfected only with KDEL-moxGFP. Fixed cells were probed for BK-CaV1.3 hetero-clusters (PLA puncta) and GFP. PLA puncta is shown in magenta. ER is shown in green. E. Comparison of BK-CaV1.3 hetero-clusters found at the ER and relative to all PLA puncta in the cell. Values are given in percentages. F. Representative images of PLA puncta and ER exit sites. Left and right are the same as in D, but cells were transfected with Sec16-GFP instead of KDEL. G. Comparison of BK-CaV1.3 hetero-clusters found at ER exit sites relative to all PLA puncta in the cell. Values are given in percentages. Data points are from n = 45 tsA-201 cells for ER, n = 21 tsA-201 cells for ERES, n = 23 INS-1 cells for ER, and n = 23 INS-1 cells for ER exit sites. Scale bar is 10 μm and 2 μm in the magnifications.

Removed in the revised version: By quantifying the proportion of BK and CaV1.3 hetero-clusters relative to total channel expression at ER exit sites, we found 28 ± 3% colocalization in tsA-201 cells and 11 ± 2% in INS-1 cells (Figure 3F). While the percentage of colocalization between hetero-clusters and the ER or ER exit sites alone cannot be directly compared to infer trafficking dynamics, these findings reinforce the conclusion that hetero-clusters reside within the ER and suggest that BK and CaV1.3 channels traffic together through the ER and exit in coordination.

BK and CaV1.3 hetero-clusters go through the Golgi

Earlier version: Channels are modified in the Golgi after synthesis (Figure 4A), so we asked whether PLA puncta formed by BK and CaV1.3 hetero-clusters can be found in the Golgi. Using the same strategy, we labeled Golgi with Gal-T-mEGFP (Figure 4B) and detected BK-CaV1.3 hetero-clusters using PLA. To confirm that the overexpressed Gal-T-mEGFP labels the Golgi specifically without altering its structure, we compared the region detected by the antibody against GFP to the region detected by a primary antibody against the Golgi protein 58K (Figure 4C). 58K is a specific peripheral protein that localizes to the cytosolic face of Golgi [31]. The overlay shows that the GFP signal labels the same region as the antibody against 58K. When assessing the presence of BK-CaV1.3 hetero-clusters in the Golgi, we found that 31 ± 5% of PLA puncta were localized at the Golgi in the overexpression system and 25 ± 4% in the endogenous cell model (Figure 4E).

Revised version: Channels are modified in the Golgi after synthesis (Figure 4A), so we asked whether PLA puncta formed by BK and CaV1.3 hetero-clusters can be found in the Golgi. Using the same strategy, we labeled Golgi with Gal-T-mEGFP (Figure 4B) and detected BK-CaV1.3 hetero-clusters using PLA. To confirm that the overexpressed Gal-T-mEGFP labels the Golgi specifically without altering its structure, we compared the region detected by the antibody against GFP to the region detected by a primary antibody against the Golgi protein 58K (Figure 4C). 58K is a specific peripheral protein that localizes to the cytosolic face of Golgi [36]. The overlay shows that the GFP signal labels the same region as the antibody against 58K. When assessing the presence of BK-CaV1.3 hetero-clusters in the Golgi, we found that 31 ± 5% of PLA puncta were localized at the Golgi in the overexpression system and 25 ± 4% in the endogenous cell model (Figure 4E).

BK and CaV1.3 hetero-clusters go through the Golgi.

A. Diagram of the hypothesis: PLA puncta detecting hetero-clusters between BK (magenta) and CaV1.3 (cyan) channels can be found at the Golgi membrane. B. Representative image of the Golgi structure with exogenous GFP in INS-1 cells. Cells were transfected with Gal-T-mEGFP. Enlargement is shown in the inset. C. Representative images of fixed cells co-stained with antibodies against Gal-T-mEGFP in green and 58K-Golgi in red D. Representative images of PLA puncta and Golgi. tsA-201 cells were transfected with BK, CaV1.3 and Gal-T-mEGFP (left), and INS-1 cells were transfected only with Gal-T-mEGFP (right). PLA puncta are shown in magenta. Golgi is shown in green. E. Scatter dot plot of percentages of BK-CaV1.3 hetero-clusters found at the Golgi relative to all PLA puncta in tsA-201 and INS1 cells. Data points are from n = 22 tsA-201 cells and n = 19 INS-1 cells. F. Representative image of PLA puncta and Golgi. tsA-201 cells were transfected with BK, CaV1.3, and Gal-T-mEGFP. Left: PLA was done against BK and 58K Golgi (magenta), and Golgi is shown in green. Right: PLA was done against CaV1.3 and 58K Golgi (magenta), and Golgi is shown in green. G. Diagram illustrating our interpretation of percentages of BK-CaV1.3 hetero-clusters found in the cell. This illustration is based on results shown in Figures 1-4. Percentages were modified to represent overlap of fluorescent signals and limited resolution. We also show that hetero-clusters found in the ER exit sites (ERES) are also accounted in the ER. Scale bars: 10 μm and 2 μm in panels B and C; 10 μm and 1 μm in panel D; 2 μm in panel F.

Earlier version: We performed controls to confirm that the formation of hetero-clusters between the two channels was not coincidental but rather the result of structural coupling. We tested the formation of PLA puncta between BK channels and the Golgi protein 58K. We also tested for CaV1.3 and 58K. We did not find PLA puncta when the proximity between the channels and 58K was probed (Figure 4F), supporting the idea that PLA puncta between BK and CaV1.3 channels found at the Golgi represent specific coupling. We selected the Golgi as a control because it represents the final stage of protein trafficking, ensuring that hetero-cluster interactions observed at this point reflect specificity maintained throughout earlier trafficking steps, including within the ER.

Revised version: We performed controls to confirm that the formation of hetero-clusters between the two channels was not coincidental but rather the result of structural coupling. We tested the formation of PLA puncta between BK channels and the Golgi protein 58K. We also tested for CaV1.3 and 58K. We did not find PLA puncta when the proximity between the channels and 58K was probed (Figure 4F), supporting the idea that PLA puncta between BK and CaV1.3 channels found at the Golgi represent specific coupling. We selected the Golgi as a control because it represents the final stage of protein trafficking, ensuring that hetero-cluster interactions observed at this point reflect specificity maintained throughout earlier trafficking steps, including within the ER. As an additional control, we tested the formation of PLA puncta between CaV1.3 channels and RyR2, a protein localized in the ER. The number of PLA puncta between CaV1.3 and RyR2 was significantly lower than the number observed between CaV1.3 and BK channels (Figure 4-Supplement 1), further supporting the specificity of BK–CaV1.3 interactions.

It is important to clarify that the percentages provided in this work cannot be added up to understand the distribution of hetero-clusters along the biosynthetic pathway of the channels. The percentage in each membrane compartment was compared to the total percentage of hetero-clusters observed in each cell for that particular experiment. Therefore, there is expected overlap between our measurements due to (1) optical resolution and (2) the effect of not comparing all the organelles in the same cell. Considering these limitations, we interpreted our results as follows: about one half of hetero-clusters are found inside the cell. The other half corresponds to hetero-clusters at the plasma membrane. From the half inside the cell, roughly one third of hetero-clusters is found in the Golgi, and another third is found in ER exit sites (Figure 4G). The remaining hetero-clusters are found in other regions of the ER and in membranes that we did not probe, such as vesicles. Finally, a key limitation of this approach is that we cannot quantify the proportion of total BK or CaV1.3 channels engaged in hetero-clusters within each compartment. The PLA method provides proximity-based detection, which reflects relative localization rather than absolute channel abundance within individual organelles.

BK mRNA and CaV1.3 mRNA colocalize

Earlier version: How is it that BK and CaV1.3 proteins come in close proximity in membranes of the ER, ER exit sites and the Golgi? To explore the origins of the initial association, we hypothesized that the two proteins are translated near each other, which could be detected as the colocalization of their mRNAs (Figure 5A and B). The experiment was designed to detect single mRNA molecules from INS-1 cells in culture. We performed multiplex in situ hybridization experiments using an RNAscope fluorescence detection kit to be able to image three mRNAs simultaneously in the same cell and acquired the images in a confocal microscope with high resolution. To rigorously assess the specificity of this potential mRNA-level organization, we used multiple internal controls. GAPDH mRNA, a highly expressed housekeeping gene with no known spatial coordination with channel mRNAs, served as a baseline control for nonspecific colocalization due to transcript abundance. To evaluate whether the spatial proximity between BK mRNA (KCNMA1) and CaV1.3 mRNA (CACNA1D) was unique to functionally coupled channels, we also tested for NaV1.7 mRNA (SCN9A), a transmembrane sodium channel expressed in INS-1 cells but not functionally associated with BK. This allowed us to determine whether the observed colocalization reflected a specific biological relationship rather than shared expression context. Finally, to test whether this proximity might extend to other calcium sources relevant to BK activation, we probed the mRNA of ryanodine receptor 2 (RyR2), another Ca2+ channel known to interact structurally with BK channels [32]. Together, these controls were chosen to distinguish specific mRNA colocalization patterns from random spatial proximity, shared subcellular distribution, or gene expression level artifacts. Figure 5-Supplementary-2 shows images completely void of fluorescent signal from INS-1 cells probed for bacterial mRNA (negative control, Figure 5-Supplementary 2A) and shows images of INS-1 cells probed for mammalian mRNAs expected to be in these cells (positive control, Figure 5-Supplementary 2B). The probes against the mRNAs of interest and tested in this work were designed by Advanced Cell Diagnostics. We confirmed the specificity of the probes by performing in situ hybridization against KCNMA1, CACNA1D, RyR2, and SCN9A in non-transfected cell lines (Figure 5-Supplementary 2C and D).

Revised version: How is it that BK and CaV1.3 proteins come in close proximity in membranes of the ER, ER exit sites and the Golgi? To explore the origins of the initial association, we hypothesized that the two proteins are translated near each other, which could be detected as the colocalization of their mRNAs (Figure 5A and B). The experiment was designed to detect single mRNA molecules from INS-1 cells in culture. We performed multiplex in situ hybridization experiments using an RNAscope fluorescence detection kit to be able to image three mRNAs simultaneously in the same cell and acquired the images in a confocal microscope with high resolution. To rigorously assess the specificity of this potential mRNA-level organization, we used multiple internal controls. GAPDH mRNA, a highly expressed housekeeping gene with no known spatial coordination with channel mRNAs, served as a baseline control for nonspecific colocalization due to transcript abundance. To evaluate whether the spatial proximity between BK mRNA (KCNMA1) and CaV1.3 mRNA (CACNA1D) was unique to functionally coupled channels, we also tested for NaV1.7 mRNA (SCN9A), a transmembrane sodium channel expressed in INS-1 cells but not functionally associated with BK. This allowed us to determine whether the observed colocalization reflected a specific biological relationship rather than shared expression context. Finally, to test whether this proximity might extend to other calcium sources relevant to BK activation, we probed the mRNA of ryanodine receptor 2 (RyR2), another Ca²⁺ channel known to interact structurally with BK channels [37]. Together, these controls were chosen to distinguish specific mRNA colocalization patterns from random spatial proximity, shared subcellular distribution, or gene expression level artifacts. Figure 5-Supplementary-2 shows images completely void of fluorescent signal from INS-1 cells probed for bacterial mRNA (negative control, Figure 5-Supplementary 2A) and shows images of INS-1 cells probed for mammalian mRNAs expected to be in these cells (positive control, Figure 5-Supplementary 2B). The probes against the mRNAs of interest and tested in this work were designed by Advanced Cell Diagnostics. We confirmed the specificity of the probes by performing in situ hybridization against KCNMA1, CACNA1D, RyR2, and SCN9A in non-transfected cell lines (Figure 5-Supplementary 2C and D).

BK mRNA (KCNMA1) and CaV1.3 mRNA (CACNA1D) colocalize.

Earlier version: A. Diagram of the hypothesis: KCNMA1 and CACNA1D mRNAs are found in close proximity to be translated in the same neighborhood. B. Images of fluorescent puncta from RNA scope experiments showing KCNMA1 mRNA in magenta, CACNA1D mRNA in cyan, and GAPDH mRNA in green. Right, magnification of three ROI. C. Comparison of mRNA density of KCNMA1, CACNA1D, and GAPDH. D. Correlation plot of mRNA abundance of KCNMA1 and CACNA1D per cell. E. Correlation plot of mRNA abundance of KCNMA1 and GAPDH per cell. F. Comparison of colocalization between KCNMA1 mRNA and mRNA from CACNA1D, GAPDH, and scrambled images of CACNA1D. Data points are from n = 67 cells. Scale bar is 10 μm and 1 μm in the magnifications.

Revised version: A. Diagram of the hypothesis: KCNMA1 and CACNA1D mRNAs are found in close proximity to be translated in the same neighborhood. B. Images of fluorescent puncta from RNA scope experiments showing KCNMA1 mRNA in magenta, CACNA1D mRNA in cyan, and GAPDH mRNA in green. Right, magnification of three ROI. C. Comparison of mRNA density of KCNMA1, CACNA1D, and GAPDH. D. Correlation plot of mRNA abundance of KCNMA1 and CACNA1D per cell. E. Correlation plot of mRNA abundance of KCNMA1 and GAPDH per cell. F. Comparison of colocalization between KCNMA1 mRNA and mRNA from CACNA1D, GAPDH, and scrambled images of CACNA1D. Data points are from n = 67 cells. Scale bar is 10 μm and 1 μm in the magnifications. Data were analyzed using ordinary one-way ANOVA with Dunnett’s multiple comparisons test.

Removed in the revised version: Illustration of RNA scope methodology.

Removed in the revised version: Steps to detect an mRNA sequence (red zipper) consisting of (1) hybridization, (2) pre-amplification, (3) and labeling. Double ZZ probes are shown in green. Dye labeling the mRNA is shown in magenta.

Removed in the revised version: RNA probe validation.

Removed in the revised version: A. Representative images of INS-1 cells subjected to mRNA probes against a bacterial gene (DapB from Bacillus subtilis). B. Representative images of INS-1 cells subjected to mRNA probes against three constitutive mammalian genes (ubiquitin C, cyclophilin B, and a polymerase II subunit). C-F show tsA-201 cells that were naive to exogenous DNA but treated with mRNA probes against the genes for (C) BK channels, (D) CaV1.3 channels, (E) Ryanodine Receptors type 2 (RyR2), and (F) NaV1.7 channels. Scale bar is 10 μm.

Earlier version: Probes against GAPDH mRNA were used as a control in the same cells that we probed for KCNMA1 and CACNA1D using the multiplex capability of this design. Transcripts were detected at different expression levels. GAPDH and CACNA1D mRNAs were more abundant (134 and 33 mRNA/100 μm2, respectively) than KCNMA1 mRNA (12 mRNA/100 μm2, Figure 5C). Interestingly, the abundance of KCNMA1 transcripts correlated more with the abundance of CACNA1D transcripts than with the abundance of GAPDH, though with a modest R2 value (Figure 5D and E). Furthermore, KCNMA1 and CACNA1D mRNA colocalized by 60 ± 4%, which was 20% more than with GAPDH mRNA (Figure 5F). As an additional control and to rule out the potential influence of difference between CACNA1D and KCNMA1 abundance, we assessed CACNA1D colocalization against randomized, computer-generated KCNMA1 mRNA signals, where localization was randomized while maintaining the same overall transcript count. The significantly lower (20%) colocalization observed in scrambled conditions compared to genuine BK-Cav1.3 mRNA interactions confirms that proximity is not an artifact of expression levels but reflects a specific spatial association. These results suggest that some fraction of mRNAs for BK and CaV1.3 channels are translated nearby, so the channel proteins potentially could be inserted into the same regions of the ER. We suggest that the newly synthesized proteins remain together during trafficking through the ER and the Golgi.

Revised version: Probes against GAPDH mRNA were used as a control in the same cells that we probed for KCNMA1 and CACNA1D using the multiplex capability of this design. Transcripts were detected at different expression levels. GAPDH and CACNA1D mRNAs were more abundant (134 and 33 mRNA/100 μm2, respectively) than KCNMA1 mRNA (12 mRNA/100 μm2, Figure 5C). Interestingly, the abundance of KCNMA1 transcripts correlated more with the abundance of CACNA1D transcripts than with the abundance of GAPDH, though with a modest R² value (Figure 5D and E). Furthermore, KCNMA1 and CACNA1D mRNA colocalized by 60 ± 4%, which was 20% more than with GAPDH mRNA (Figure 5F). As an additional control and to rule out the potential influence of difference between CACNA1D and KCNMA1 abundance, we assessed CACNA1D colocalization against randomized, computer-generated KCNMA1 mRNA signals, where localization was randomized while maintaining the same overall transcript count. The significantly lower (20%) colocalization observed in scrambled conditions compared to genuine BK-Cav1.3 mRNA interactions confirms that proximity is not an artifact of expression levels but reflects a specific spatial association. These results suggest that some fraction of mRNAs for BK and CaV1.3 channels are translated nearby, so the channel proteins potentially could be inserted into the same regions of the ER. We suggest that the newly synthesized proteins remain together during trafficking through the ER and the Golgi.

Earlier version: Next, we compared the proximity between KCNMA1 and SCN9A. We detected SCN9A in the same cells where KCNMA1 was found (Figure 6A). The colocalization between KCNMA1 and SCN9A was only 18 ± 2% (Figure 6C), which was less than what was observed with KCNMA1 and CACNA1D (60%, Figure 5F). This result shows that KCNMA1 tends to localize closer to CACNA1D than to SCN9A or GAPDH, supporting the specificity of KCNMA1 and CACNA1D mRNA association.

Revised version: Next, we compared the proximity between KCNMA1 and SCN9A. We detected SCN9A in the same cells where KCNMA1 was found (Figure 6A). The colocalization between KCNMA1 and SCN9A was only 18 ± 2% (Figure 6C), which was less than what was observed with KCNMA1 and CACNA1D (60%, Figure 5F). notably, We observed comparable levels of co-localization between KCNMA1 and SCN9A transcripts in both directions (data not shown), with no statistically significant difference. These findings support the specificity of mRNA co-localization and that KCNMA1 tends to localize closer to CACNA1D than to SCN9A or GAPDH, supporting the specificity of KCNMA1 and CACNA1D mRNA association.

BK mRNA (KCNMA1) and RyR-2 mRNA (RyR2) colocalize.

Earlier version: A. Representative confocal images of KCNMA1 and NaV1.7 (SCN9A) mRNA. B. Representative images of KCNMA1 and RyR2 mRNA. C. Comparison of the colocalization between KCNMA1 mRNA and mRNA from RyR2, SCN9A, and scrambled images of KCNMA1. Data points are from n =67 cells. One way ANOVA. Scale bar is 10 μm and 1 μm in the magnifications.

Revised version: A. Representative confocal images of KCNMA1 and NaV1.7 (SCN9A) mRNA. B. Representative images of KCNMA1 and RyR2 mRNA. C. Comparison of the colocalization between KCNMA1 mRNA and mRNA from RyR2, SCN9A, and scrambled images of KCNMA1. Data points are from n =67 cells. One way ANOVA. Scale bar is 10 μm and 1 μm in the magnifications. Data were analyzed using ordinary one-way ANOVA with Dunnett’s multiple comparisons test.

Earlier version: With the goal of understanding if this concept could apply to other channels, we used the same approach to test a second protein known to provide calcium for BK channel opening. Similar to the coupling between BK and CaV1.3 channels, RyR2 also can be structurally coupled to BK channels [32]. Figure 6B shows high resolution images of single molecule in situ hybridization for RyR2 probed together with KCNMA1 in INS-1 cells. In the same population of cells, KCNMA1 and RyR2 colocalization was 27 ± 3%, which is 1.5 times that with SCN9A mRNA. We also assessed RyR2 colocalization with randomized, computer-generated KCNMA1 mRNA signals. We found that colocalization between randomized KCNMA1 and genuine RyR2 was 87% less colocalized compared to the analysis of genuine KCNMA1 and genuine RyR2 (Figure 6C), suggesting that KCNMA1 not only colocalizes with CACNA1D but also with mRNA of other known calcium sources.

Revised version: With the goal of understanding if this concept could apply to other channels, we used the same approach to test a second protein known to provide calcium for BK channel opening. Similar to the coupling between BK and CaV1.3 channels, RyR2 also can be structurally coupled to BK channels [37]. Figure 6B shows high resolution images of single molecule in situ hybridization for RyR2 probed together with KCNMA1 in INS-1 cells. In the same population of cells, KCNMA1 and RyR2 colocalization was 27 ± 3%, which is 1.5 times that with SCN9A mRNA. We also assessed RyR2 colocalization with randomized, computer-generated KCNMA1 mRNA signals. We found that colocalization between randomized KCNMA1 and genuine RyR2 was 87% less colocalized compared to the analysis of genuine KCNMA1 and genuine RyR2 (Figure 6C), suggesting that KCNMA1 not only colocalizes with CACNA1D but also with mRNA of other known calcium sources.

Earlier version: To further investigate whether KCNMA1 and CACNA1D are localized in regions of active translation (Figure 7A), we performed RNAscope targeting KCNMA1 and CACNA1D alongside immunostaining for BK protein. This strategy enabled us to visualize transcript-protein colocalization in INS-1 cells with subcellular resolution. By directly evaluating sites of active BK translation, we aimed to determine whether newly synthesized BK protein colocalized with CACNA1D mRNA signals (Figure 7A). Confocal imaging revealed distinct micro-translational complex where KCNMA1 mRNA puncta overlapped with BK protein signals and were located adjacent to CACNA1D mRNA (Figure 7B). Quantitative analysis showed that 71 ± 3% of all KCNMA1 colocalized with BK protein signal which means that they are in active translation. Interestingly, 69 ± 3% of the KCNMA1 in active translation colocalized with CACNA1D (Figure 7C), supporting the existence of functional micro-translational complexes between BK and CaV1.3 channels.

Revised version: To further investigate whether KCNMA1 and CACNA1D are localized in regions of active translation (Figure 7A), we performed RNAscope targeting KCNMA1 and CACNA1D alongside immunostaining for BK protein. This strategy enabled us to visualize transcript-protein colocalization in INS-1 cells with subcellular resolution.

BK mRNA (KCNMA1) and CaV1.3 mRNA (CACNA1D) colocalize in micro-translational complexes.

A. Diagram of the hypothesis: KCNMA1 mRNAs are found in micro-translational complexes. B. Representative images of KCNMA1 mRNA in magenta, CACNA1D mRNA in cyan, and BK protein in green. C. Comparison of the frequency of colocalization KCNMA1 mRNA in active translation and in micro-translational complexes. Data points are from n = 57 cells. One way ANOVA was used as statistical analysis. Scale bar is 10 μm and 1 μm in the magnifications.

Added in the revised version: By directly evaluating sites of active BK translation, we aimed to determine whether newly synthesized BK protein colocalized with CACNA1D mRNA signals (Figure 7A). Confocal imaging revealed distinct micro-translational complex where KCNMA1 mRNA puncta overlapped with BK protein signals and were located adjacent to CACNA1D mRNA (Figure 7B). Quantitative analysis showed that 71 ± 3% of all KCNMA1 colocalized with BK protein signal which means that they are in active translation. Interestingly, 69 ± 3% of the KCNMA1 in active translation colocalized with CACNA1D (Figure 7C), supporting the existence of functional micro-translational complexes between BK and CaV1.3 channels.

BK and CaV1.3 channels form hetero-clusters at the plasma membrane of INS-1 cells

Added in the revised version: We previously showed the organization of BK and CaV1.3 channels in hetero-clusters in tsA-201 cells and neurons, including hippocampal and sympathetic motor neurons [15]. Our present study utilizes INS-1 cells. We detected single microscopy and determined their degree of colocalization. Figure 8A shows a representative image of the localizations of antibodies against BK and CaV1.3 channels. Maps were rendered at 5 nm, the pixel size for the images presented. BK-positive pixels formed multi-pixel aggregates, which has been interpreted as homo-clusters of molecules in INS-1 cells. It is important to note that this technique does not allow us to distinguish between labeling of four BK α-subunits within a tetramer and labeling of multiple BK channels in a cluster. Hence, particles smaller than ∼1680 nm² may represent either a single tetramer or a cluster. This limitation applies to Figures 8C and 9D and does not affect measurements of BK– CaV1.3 proximity. CaV1.3 channels also formed homo-clusters. When looking at the formation of BK and CaV1.3 hetero-clusters, no fixed geometry or stoichiometry was observed. On the contrary, maps showed a distribution of distances between BK and CaV1.3 detected pixels. In some cases, BK and CaV1.3 channels were perfectly overlapping, as if the detection of the antibodies could not be separated by the 20 nm resolution. In other cases, BK and CaV1.3 channels were adjacent (less than 5 nm), and in other cases, BK homo-clusters were surrounded by a variable number of CaV1.3 channels.

Removed in the revised version: We previously showed the organization of BK and CaV1.3 channels in hetero-clusters in tsA-201 cells and neurons, including hippocampal and sympathetic motor neurons [33]. Our present study utilizes INS-1 cells. We detected single molecules from BK and CaV1.3 channels in INS-1 cells using super-resolution microscopy and determined their degree of colocalization. Figure 8A shows a representative image of the localizations of antibodies against BK and CaV1.3 channels.

Formation of BK and CaV1.3 hetero-clusters in INS-1 cells.

Earlier version: A. Representative localization map of antibodies against BK (magenta) and CaV1.3 (cyan) channels. Magnifications are shown in the insets on the right. B. Scatter dot plot of homo-cluster densities of BK and CaV1.3 channels in INS-1 cells. C. Cumulative frequency distributions of homo-cluster sizes of BK and CaV1.3 channels. Inset compares median size of BK and CaV1.3 homo-clusters. D. Comparison of colocalization between BK and CaV1.3 and between scrambled BK and CaV1.3. Data points are from n = 13 cells. Scale bar is 5 μm and 300 nm in the magnifications.

Revised version: A. Representative localization map of antibodies against BK (magenta) and CaV1.3 (cyan) channels. Magnifications are shown in the insets on the right. B. Scatter dot plot of homo-cluster densities of BK and CaV1.3 channels in INS-1 cells. C. Cumulative frequency distributions of homo-cluster sizes of BK and CaV1.3 channels. Inset compares median size of BK and CaV1.3 homo-clusters. D. Comparison of colocalization between BK and CaV1.3 and between scrambled BK and CaV1.3. Data points are from n = 13 cells. Scale bar is 5 μm and 300 nm in the magnifications. Data was analyzed using paired t-tests to evaluate significance.

Added in the revised version: Hetero-clusters of BK and CaV1.3 channels are detected at the plasma membrane soon after their expression begins.

Added in the revised version: A-C. Representative localization maps of antibodies against BK and CaV1.3 channels. (A) 18 hours, (B) 24 hours, or (C) 48 hours after DNA transfection into cells. Enlargements are shown in insets. D. Cumulative frequency distributions of BK homo-cluster size at 18, 24, and 48 hours. Inset compares median BK homo-cluster areas. E. Cumulative frequency distributions of CaV1.3 homo-clusters at 24 and 48 hours. Inset compares median CaV1.3 cluster areas. CaV1.3 clusters are not present at the 18-hour time point. F. Comparison of colocalization plots between BK and CaV1.3 channels at 24 and 48-hour time points. Data points are from n = 10 cells. Scale bar is 10 μm and 300 nm in enlargements. Statistical significance was assessed using ordinary one-way ANOVA followed by Dunnett’s multiple comparisons test and unpaired t-test in panel D and E respectively.

Removed in the revised version: Maps were rendered at 5 nm, the pixel size for the images presented. BK-positive pixels formed multi-pixel aggregates, which has been interpreted as homo-clusters of BK channels. CaV1.3 channels also formed homo-clusters. When looking at the formation of BK and CaV1.3 hetero-clusters, no fixed geometry or stoichiometry was observed. On the contrary, maps showed a distribution of distances between BK and CaV1.3 detected pixels. In some cases, BK and CaV1.3 channels were perfectly overlapping, as if the detection of the antibodies could not be separated by the 20 nm resolution. In other cases, BK and CaV1.3 channels were adjacent (less than 5 nm), and in other cases, BK homo-clusters were surrounded by a variable number of CaV1.3 channels.

Earlier version: To provide a more quantitative description of the maps, we measured cluster density, cluster size, and colocalization. We used the particle analysis tool of the software ImageJ for these measurements. In this analysis, a particle for represents positive pixels irrespective of the number of pixels composed. As a note, a particle of one pixel would likely to be a single channel, whereas a particle composed of many pixels would likely to be a cluster of channels. INS-1 cells showed 4 times as many CaV1.3 particles as BK particles (Figure 8B). Figure 8C shows the frequency distribution of particle size, where the median area is 1691 nm2 for BK and 975 nm2 for CaV1.3. We also found 7.5% colocalization between BK and CaV1.3 particles (Figure 8D), which was higher than the colocalization calculated using scrambled images.

Revised version: To provide a more quantitative description of the maps, we measured cluster density, cluster size, and colocalization. We used the particle analysis tool of the software ImageJ for these measurements. In this analysis, a particle for represents positive pixels irrespective of the number of pixels composed. INS-1 cells showed 4 times as many CaV1.3 particles as BK particles (Figure 8B). Figure 8C shows the frequency distribution of particle size, where the median area is 1691 nm2 for BK and 975 nm2 for CaV1.3. The colocalization between BK and CaV1.3 clusters (Figure 8D) was 7.5%, a value higher than that obtained from scrambled image controls. Notably, 37 ± 3% BK channels are at 200 nm or less from one or more CaV1.3, while 15 ± 2% of CaV1.3 channels are at 200 nm or less from one or more BK (Figure 8-supplementary 1 Panel A). Furthermore, the distribution of the nearest distance between BK and CaV1.3 in INS-1 cells (Figure 8-supplementary 1 Panel B, magenta) shows a large proportion of BK channels (around 40%) at <50 nm from any CaV1.3 channels. Together, these results suggest that INS-1 cells also exhibit nanodomains containing BK-CaV1.3 hetero-clusters at the plasma membrane.

Hetero-clusters of BK and CaV1.3 channels are detected at the plasma membrane soon after their expression

Earlier version: In light of our results, our current model is that BK and CaV1.3 hetero-clusters form prior to their insertion into the plasma membrane. One prediction based on this model is that channels inserted in the plasma membrane would appear as hetero-clusters already at early time points after the start of their synthesis. To test this prediction, we transfected tsA-201 cells with BK and CaV1.3 channels and performed a chase experiment to detect their presence at the plasma membrane using super-resolution microscopy. We measured particle size, density, and colocalization at 18, 24, and 48 hours after transfection. Notably, although the channels were transfected simultaneously, BK particles were detected in the plasma membrane 18 hours after transfection, whereas CaV1.3 particles were detected 24 hours after transfection (Figure 9A-C). The distribution of BK and CaV1.3 particle size did not change with the time after transfection (Figure 9D and E), in agreement with Sato et al., 2019 [24].

Revised version: In light of our results, our current model is that BK and CaV1.3 hetero-clusters form prior to their insertion into the plasma membrane. One prediction based on this model is that channels inserted in the plasma membrane would appear as hetero-clusters already at early time points after the start of their synthesis. To test this prediction, we transfected tsA-201 cells with BK and CaV1.3 channels and performed a chase experiment to detect their presence at the plasma membrane using super-resolution microscopy. We measured particle size, density, and colocalization at 18, 24, and 48 hours after transfection. Notably, although the channels were transfected simultaneously, BK particles were detected in the plasma membrane 18 hours after transfection, whereas CaV1.3 particles were detected 24 hours after transfection (Figure 9A-C). The distribution of BK and CaV1.3 particle size did not change with the time after transfection (Figure 9D and E), in agreement with Sato et al., 2019 [30].

Removed in the revised version: Hetero-clusters of BK and CaV1.3 channels are detected at the plasma membrane soon after their expression begins.

Removed in the revised version: A-C. Representative localization maps of antibodies against BK and CaV1.3 channels. (A) 18 hours, (B) 24 hours, or (C) 48 hours after DNA transfection into cells. Enlargements are shown in insets. D. Cumulative frequency distributions of BK homo-cluster size at 18, 24, and 48 hours. Inset compares median BK homo-cluster areas. E. Cumulative frequency distributions of CaV1.3 homo-clusters at 24 and 48 hours. Inset compares median CaV1.3 cluster areas. CaV1.3 clusters are not present at the 18-hour time point. F. Comparison of colocalization plots between BK and CaV1.3 channels at 24 and 48-hour time points. Data points are from n = 10 cells. Scale bar is 10 μm and 300 nm in enlargements.

In support of our prediction, we found particles of CaV1.3 near BK particles as soon as the expression of CaV1.3 channels was detected in the plasma membrane (24 hours after transfection, enlarged Figure 9B). To test this hypothesis further, we compared plots of localization preference at 24 and 48 hours. These plots were constructed by measuring the percentage of CaV1.3 area occupying concentric regions of 20 nm width around BK particles. Values were normalized to the percentage of colocalization found at 200 nm from BK. Figure 9F shows that the localization preference plots are identical at 24 and 48 hours, consistent with the hypothesis that both channels are inserted together into the plasma membrane already as hetero-clusters.

Discussion

Intracellular Assembly of BK and CaV1.3 Channels

Earlier version: Previous work from our group has revealed the organization of hetero-clusters of large-conductance calcium-activated potassium (BK) channels and voltage-gated calcium channels, particularly CaV1.3, within nanodomains. This spatial organization is crucial for the functional coupling that enables BK channels to respond promptly to calcium influx, thereby modulating cellular excitability [15]. While physical details of BK-CaV1.3 interactions are not fully elucidated, the evidence suggests that BK channels interact with both the CaV α1 subunit and its auxiliary subunits [34, 35]. Our findings highlight the intracellular assembly of BK-CaV1.3 hetero-clusters, though limitations in resolution and organelle-specific analysis prevent precise quantification of the proportion of intracellular complexes that ultimately persist on the cell surface. While our data confirms that hetero-clusters form before reaching the plasma membrane, it remains unclear whether all intracellular hetero-clusters transition intact to the membrane or undergo rearrangement or disassembly upon insertion. Future studies utilizing live cell tracking and high resolution imaging will be valuable in elucidating the fate and stability of these complexes after membrane insertion.

Revised version: Previous work from our group has revealed the organization of hetero-clusters of large-conductance calcium-activated potassium (BK) channels and voltage-gated calcium channels, particularly CaV1.3, within nanodomains. This spatial organization is crucial for the functional coupling that enables BK channels to respond promptly to calcium influx, thereby modulating cellular excitability. While physical details of BK-CaV1.3 interactions are not fully elucidated, the evidence suggests that BK channels interact with both the CaV α1 subunit and its auxiliary subunits [38, 39].

Added in the revised version: Our findings highlight the intracellular assembly of BK-CaV1.3 hetero-clusters, though limitations in resolution and organelle-specific analysis prevent precise quantification of the proportion of intracellular complexes that ultimately persist on the cell surface. While our data confirms that hetero-clusters form before reaching the plasma membrane, it remains unclear whether all intracellular hetero-clusters transition intact to the membrane or undergo rearrangement or disassembly upon insertion. Future studies utilizing live cell tracking and high resolution imaging will be valuable in elucidating the fate and stability of these complexes after membrane insertion.

Earlier version: A stochastic model of ion channel homo-cluster formation in the plasma membrane proposes that random yet probabilistic interactions between ion channels contribute to their formation [24]. These clusters are stabilized and organized within specific membrane regions through biophysical mechanisms such as membrane curvature. Recent findings on spatial organization of G-protein coupled receptors provide additional support for this framework, demonstrating that the coupling of membrane proteins to the curvature of the plasma membrane acts as a driving force for their clustering into functional domains [36]. Yet, the mechanisms regulating the spatial organization of hetero-clusters are less known.

Revised version: A stochastic model of ion channel homo-cluster formation in the plasma membrane proposes that random yet probabilistic interactions between ion channels contribute to their formation [30]. These clusters are stabilized and organized within specific membrane regions through biophysical mechanisms such as membrane curvature. Recent findings on spatial organization of G-protein coupled receptors provide additional support for this framework, demonstrating that the coupling of membrane proteins to the curvature of the plasma membrane acts as a driving force for their clustering into functional domains [40]. Yet, the mechanisms regulating the spatial organization of hetero-clusters are less known.

Colocalization and Trafficking Dynamics

Earlier version: The colocalization of BK and CaV1.3 channels in the ER and at ER exit sites before reaching the Golgi suggests a coordinated trafficking mechanism that facilitates the formation of multi-channel complexes crucial for calcium signaling and membrane excitability [37, 38]. Given the distinct roles of these compartments, colocalization at the ER and ER exit sites may reflect transient proximity rather than stable interactions. Their presence in the Golgi further suggests that post-translational modifications and additional assembly steps occur before plasma membrane transport, providing further insight into hetero-cluster maturation and sorting events. By examining BK-CaV1.3 hetero-cluster distribution across these trafficking compartments, we ensure that observed colocalization patterns are considered within a broader framework of intracellular transport mechanisms [39]. Previous studies indicate that ER exit sites exhibit variability in cargo retention and sorting efficiency [40], emphasizing the need for careful evaluation of colocalization data. Accounting for these complexities allows for a robust assessment of signaling complexes formation and trafficking pathways.

Revised version: The colocalization of BK and CaV1.3 channels in the ER and at ER exit sites before reaching the Golgi suggests a coordinated trafficking mechanism that facilitates the formation of multi-channel complexes crucial for calcium signaling and membrane excitability [41, 42]. Given the distinct roles of these compartments, colocalization at the ER and ER exit sites may reflect transient proximity rather than stable interactions. Their presence in the Golgi further suggests that post-translational modifications and additional assembly steps occur before plasma membrane transport, providing further insight into hetero-cluster maturation and sorting events. By examining BK-CaV1.3 hetero-cluster distribution across these trafficking compartments, we ensure that observed colocalization patterns are considered within a broader framework of intracellular transport mechanisms [43]. Previous studies indicate that ER exit sites exhibit variability in cargo retention and sorting efficiency [44], emphasizing the need for careful evaluation of colocalization data. Accounting for these complexities allows for a robust assessment of signaling complexes formation and trafficking pathways.

BK Surface Expression and Independent Trafficking Pathways

Earlier version: BK surface expression in the absence of CaV1.3 indicates that its trafficking does not strictly rely on CaV1.3-mediated interactions. Since BK channels can be activated by multiple calcium sources, their presence in intracellular compartments suggests that their surface expression is governed by intrinsic trafficking mechanisms rather than direct calcium-dependent regulation. While some BK and CaV1.3 hetero-clusters assemble into signaling complexes intracellularly, other BK channels follow independent trafficking pathways, demonstrating that complex formation is not obligatory for all BK channels. Differences in their transport kinetics further reinforce the idea that their intracellular trafficking is regulated through distinct mechanisms. Studies have shown that BK channels can traffic independently of CaV1.3, relying on alternative calcium sources for activation [13, 41]. Additionally, CaV1.3 exhibits slower synthesis and trafficking kinetics than BK, emphasizing that their intracellular transport may not always be coordinated. These findings suggest that BK and CaV1.3 exhibit both independent and coordinated trafficking behaviors, influencing their spatial organization and functional interactions.

Revised version: BK surface expression in the absence of CaV1.3 indicates that its trafficking does not strictly rely on CaV1.3-mediated interactions. Since BK channels can be activated by multiple calcium sources, their presence in intracellular compartments suggests that their surface expression is governed by intrinsic trafficking mechanisms rather than direct calcium-dependent regulation. While some BK and CaV1.3 hetero-clusters assemble into signaling complexes intracellularly, other BK channels follow independent trafficking pathways, demonstrating that complex formation is not obligatory for all BK channels. Differences in their transport kinetics further reinforce the idea that their intracellular trafficking is regulated through distinct mechanisms. Studies have shown that BK channels can traffic independently of CaV1.3, relying on alternative calcium sources for activation [13, 45]. Additionally, CaV1.3 exhibits slower synthesis and trafficking kinetics than BK, emphasizing that their intracellular transport may not always be coordinated. These findings suggest that BK and CaV1.3 exhibit both independent and coordinated trafficking behaviors, influencing their spatial organization and functional interactions. Future experiments using inducible constructs to precisely control transcription timing will enable more precise quantification of hetero-cluster formation in the ER compartment prior to plasma membrane insertion and reduce the variability introduced by differences in expression timing after plasmid transfection.

mRNA Colocalization and Protein Trafficking

Earlier version: The colocalization of mRNAs encoding BK and CaV1.3 channels suggests a coordinated translation mechanism that facilitates their proximal synthesis and subsequent assembly into functional complexes. As a mechanism, mRNA localization enhances protein enrichment at functional sites, coordinating with translational control to position ion channels at specific subcellular domains; ion channel positioning is especially critical for neurons and cardiomyocytes [42, 43]. By synthesizing these channels in close proximity, channels can be efficiently assembled into functional units, making this process crucial for precise targeting and trafficking of ion channels to specific subcellular domains, thereby ensuring proper cellular function and efficient signal transduction [42, 44-46]. Disruption in mRNA localization or protein trafficking can lead to ion channel mislocalization, resulting in altered cellular function and disease [47, 48]. In cardiac cells, precise trafficking of ion channels to specific membrane subdomains is crucial for maintaining normal electrical and mechanical functions, and its disruption contributes to heart disease and arrhythmias [32, 49, 50] Similarly, in neurons, the localization and translation of mRNA at synaptic sites are essential for synaptic plasticity, and disturbances can lead to neurological disorders [51, 52].

Revised version: The colocalization of mRNAs encoding BK and CaV1.3 channels suggests a coordinated translation mechanism that facilitates their proximal synthesis and subsequent assembly into functional complexes. As a mechanism, mRNA localization enhances protein enrichment at functional sites, coordinating with translational control to position ion channels at specific subcellular domains; ion channel positioning is especially critical for neurons and cardiomyocytes [46, 47]. By synthesizing these channels in close proximity, channels can be efficiently assembled into functional units, making this process crucial for precise targeting and trafficking of ion channels to specific subcellular domains, thereby ensuring proper cellular function and efficient signal transduction [48–50]. Disruption in mRNA localization or protein trafficking can lead to ion channel mislocalization, resulting in altered cellular function and disease [51, 52]. In cardiac cells, precise trafficking of ion channels to specific membrane subdomains is crucial for maintaining normal electrical and mechanical functions, and its disruption contributes to heart disease and arrhythmias [53–55] Similarly, in neurons, the localization and translation of mRNA at synaptic sites are essential for synaptic plasticity, and disturbances can lead to neurological disorders [56, 57].

Co-translational Regulation and Functional Coordination

Earlier version: The co-translational association of transcripts encoding BK and CaV1.3 ion channels, as well as other known calcium sources like RyR2, may serve as a mechanism to ensure the precise stoichiometry and assembly of functional channel complexes. This precise assembly is also critical for spatially organizing RyR2 and BK channels, as shown in airway myocytes, where the proximity of these channels underlies the efficient activation of BK channels by localized Ca2+ sparks [32]. Co-translational association of transcripts encoding KV1.3 channels was the first example in which the interaction of nascent KV1.3 N-termini facilitates proper tertiary and quaternary structure required for oligomerization [53, 54]. Similarly, co-translational heteromeric association of hERG1a and hERG1b subunits ensures that cardiac IKr currents exhibit the appropriate biophysical properties and magnitude necessary for normal ventricular action potential shaping [45]. The precise mechanism by which BK transcripts are associated with the transcripts of calcium sources like CaV1.3 and RyR2 channels remain to be elucidated. While it is conceivable that complementary base pairing or tertiary structural interactions play a role, RNA binding proteins (RBPs) are likely key mediators of these associations.

Revised version: The potential co-translational association of transcripts encoding BK and CaV1.3 ion channels, as well as other known calcium sources like RyR2, may serve as a mechanism to ensure the precise stoichiometry and assembly of functional channel complexes. It is important to note that while our data suggest mRNA coordination, additional experiments are required to directly assess co-translation. In parallel, the spatial organization of BK channels with other calcium sources such as RyR2—particularly in airway myocytes—has been shown to facilitate efficient BK activation by localized Ca²⁺ sparks [38]. Co-translational association of transcripts encoding KV1.3 channels was the first example in which the interaction of nascent KV1.3 N-termini facilitates proper tertiary and quaternary structure required for oligomerization [58, 59]. Similarly, co-translational heteromeric association of hERG1a and hERG1b subunits ensures that cardiac IKr currents exhibit the appropriate biophysical properties and magnitude necessary for normal ventricular action potential shaping [51]. The precise mechanism by which BK transcripts are associated with the transcripts of calcium sources like CaV1.3 and RyR2 channels remain to be elucidated. While it is conceivable that complementary base pairing or tertiary structural interactions play a role, RNA binding proteins (RBPs) are likely key mediators of these associations.

Conclusion and Future Directions

This study provides novel insights into the organization of BK and CaV1.3 channels in hetero-clusters, emphasizing their assembly within the ER, at ER exit sites, and within the Golgi. Our findings suggest that BK and CaV1.3 channels begin assembling intracellularly before reaching the plasma membrane, shaping their spatial organization, and potentially facilitating functional coupling. While this suggests a coordinated process that may contribute to functional coupling, further investigation is needed to determine the extent to which these hetero-clusters persist upon membrane insertion. While our study advances the understanding of BK and CaV1.3 hetero-cluster assembly, several key questions remain unanswered. What molecular machinery drives this colocalization at the mRNA and protein level? How do disruptions to complex assembly contribute to channelopathies and related diseases? Additionally, a deeper investigation into the role of RNA binding proteins in facilitating transcript association and localized translation is warranted.